Neuro-evolutionary evidence for a universal fractal primate brain shape

Author Response:

We thank all reviewers for their comments and effort to improve our paper. We appreciate that the writing can be clarified overall, and some sections need more elaboration. We will provide these in the next revision within the coming months. Particularly, we will focus on some common themes identified by all reviewers:

1. We will clarify that the coarse-grained brain surfaces are an output of our algorithm alone and not to be directly/naively likened to actual brain surfaces, e.g. in terms of the location or shape of the folds. Our analysis purely focuses on the likeliness in terms of whole-brain morphometrics between actual brains and coarse-grained brains. Specifically on the point of “thickening” of the brain: this is anatomically well-founded, as less folded brains have a “thicker” cortex than more folded brains, when they are all normalised to the same size. This is fundamentally why the universal scaling law also applies to these coarse-grained brains. We will provide more detail to highlight this.

2. We will clarify the motivation behind our coarse-graining procedure better: mathematically, this is directly inspired by box-counting algorithms in fractal geometry; but this algorithm also has elegant parallels with other algorithms which we will highlight.

3. The age effects are demonstrated here in a small sample as a proof-of-principle, but we will update our latest results using ~100 subjects from the CamCAN data demonstrating the same effect. We have additionally described and verified these age effects in more detail in a separate preprint (https://arxiv.org/abs/2311.13501) with ~1500 subjects, and additionally showed that scale-dependent metrics substantially improve understanding and applications such as brain age prediction.

4. We have independently also received the feedback that we need to clarify how our method interacts with different resolution of the original MRI. We will add this as a new set of results, demonstrating that the MRI acquisition resolution (within a reasonable range) has a very small effect, as our method takes the reconstructed surfaces as a starting point.

5. We agree that it may be confusing to emphasise a constant K in the first set of results across species, and then later highlight a changing K in the human ageing results. We will clarify that in the first set of results, we find a “constant” K relative to a changing S: The range in K across melted primate brains is approx 0.1, whereas in S it is over 1.2. In other words, S changes are an order of magnitude higher than K changes. Hence, we described K as “constant” relative to S. Nevertheless, K shows subtle changes within individuals, which is what we are describing in the human ageing results. These changes are within the range of K values described in the across species results.

6. Finally, we will also make sure to summarise our specific contributions beyond existing work:

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The cerebral cortex displays a bewildering diversity of shapes and sizes across and within species. Despite this diversity, we present a parsimonious and universal multi-scale description of primate cortices in full agreement with empirical data, by expressing cortical shapes explicitly as hierarchical compositions of folds across spatial scales. As one gradually merges neighbouring folds of increasing size, the cortices of 11 primate species follow a common scale-free morphometric trajectory, that also overlaps with over 70 other mammalian species. Our results indicate that all cerebral cortices are approximations of the same archetypal fractal shape. Importantly, this new understanding enables a more precise quantification of neurobiological processes as a function of scale. Specifically, we demonstrate a scale-dependent effect of ageing, with a nearly two-fold increase in effect size at a spatial scale of approximately 2 mm compared to standard morphological analysis, suggesting our methods may generate superior biomarkers for a range of conditions in future.

Revised version: The cerebral cortex displays a bewildering diversity of shapes and sizes across and within species. Despite this diversity, we present a universal multi-scale description of primate cortices. We show that all cortical shapes can be described as a set of nested folds of different sizes. As neighbouring folds are gradually merged, the cortices of 11 primate species follow a common scale-free morphometric trajectory, that also overlaps with over 70 other mammalian species. Our results indicate that all cerebral cortices are approximations of the same archetypal fractal shape with a fractal dimension of df = 2.5. Importantly, this new understanding enables a more precise quantification of brain morphology as a function of scale. To demonstrate the importance of this new understanding, we show a scale-dependent effect of ageing on brain morphology. We observe a more than four-fold increase in effect size (from 2 standard deviations to 8 standard deviations) at a spatial scale of approximately 2 mm compared to standard morphological analyses. Our new understanding may therefore generate superior biomarkers for a range of conditions in the future.

Earlier version: 1 Introduction

Revised version: Introduction

Earlier version: The morphological complexity of the mammalian cerebral cortex has fascinated scientists for generations, with cortices across and within species exhibiting a large diversity of shapes and sizes. Such diversity is not arbitrary, however. The mammalian brain folds into stereotypical, hierarchically-organized structures such as lobes and major gyri. In fact, qualitative and quantitative regularities in cortical scaling have often been suggested and observed [1, 2, 3, 4]. More specifically, through modelling the mechanism of cortical folding from a statistical physics approach, we have previously derived a theoretical scaling law relating pial surface area At, exposed surface area Ae, and average cortical thickness T: AtT12=kAe54. (1)

Revised version: The morphological complexity of the mammalian cerebral cortex has fascinated scientists for generations, with cortices across and within species exhibiting a large diversity of shapes and sizes. Such diversity is not arbitrary, however. The mammalian brain folds into stereotypical, hierarchically-organized structures such as lobes and major gyri. In fact, qualitative and quantitative regularities in cortical scaling have often been suggested and observed [1, 2, 3, 4]. More specifically, through modelling the mechanism of cortical folding from a statistical physics approach, we have previously derived a theoretical scaling law relating pial surface area At, exposed surface area Ae1, and average cortical thickness T:

Earlier version: This relationship was shown to be valid across mammalian species [5] and within the human species [6], as well as to the structures and substructures of individual brains [7, 8]. Notably, across all these cases the dimensionless offset k is shown to be near invariant. However, this universality, presumed to be arising from universal physical principles and evolutionarily conserved biomechanics, says little about what that cortical shape actually is, beyond a constraint binding three morphometric parameters. In this paper, we take Eqn. (1) as an empirical starting point to create a new, elegant and hierarchical way of expressing cortical shape. Specifically, we introduce a coarse-graining procedure that will render cortical shapes at different spatial scales, or resolutions. We show that the shapes of primate cerebral cortices at each spatial scale can be understood as approximations of the same universal self-similar archetypal form.

Revised version: This scaling law, relating powers of cortical thickness and surface area metrics, was shown to be valid across mammalian species [5] and within the human species [6], as well as to the structures and substructures of individual brains [7, 8]. Notably, across all these cases the dimensionless offset k is shown to be near invariant. However, this universality, presumed to be arising from universal physical principles and evolutionarily conserved biomechanics, says little about what that cortical shape actually is, beyond a constraint binding three morphometric parameters. In this paper, we take Eqn. (1) as an empirical starting point to create a new and hierarchical way of expressing cortical shape. Specifically, we introduce a coarse-graining procedure that renders cortical shapes at different spatial scales, or resolutions. We show that the shapes of primate cerebral cortices at each spatial scale can be understood as approximations of the same universal self-similar archetypal form.

Earlier version: Besides revealing a symmetry in nature hidden under much apparent complexity, our results indicates a conservation of some quantities across evolution. We will show that these results further provide us with a new and powerful tool to express and analyze cortical morphology. As an example, we will calculate the effects of human aging across spatial scales and show that the effects are highly scale-dependent.

Revised version: Besides revealing a symmetry in nature hidden under much apparent complexity, our results indicate a conservation of some quantities across evolution. We will show that these results further provide us with a new and powerful tool to express and analyse cortical morphology. As an example, we will calculate the effects of human ageing across spatial scales and show that the effects are highly scale-dependent.

Earlier version: 2 Coarse-graining procedure

Revised version: 2 Coarse-graining Method

Removed in the revised version: The universal scaling law (Eqn. 1) can be rewritten in a suggestive way AtA0=(AeA0)1.25 (2)

Earlier version: where the A0=T2k4 is a fundamental area element that defines the threshold between gyrencephaly and lissencephaly when At = Ae = A0. For a constant k the value of A0 is a multiple of T2, indicating that cortical thickness determines the size of the smallest possible gyri and sulci.

Revised version: The universal scaling law (Eqn. 1) can be rewritten in a suggestive way A t A 0 = A e A 0 1.25 , (2) where the A 0 = T 2 k 2 is a fundamental area element that defines the threshold between gyrencephaly (folded cortex) and lissencephaly (smooth cortex) when At = Ae = A0. For a constant k the value of A0 is a multiple of T2, indicating that cortical thickness determines the size of the smallest possible gyri and sulci.

Earlier version: This re-writing highlights a new perspective, or interpretation of the scaling law: it now suggests a relationship between intrinsic and extrinsic measures of cortical size1: At and Ae respectively, measured in units of A0. This is reminiscent of fractal scaling [9], where a complex shape reveals ever smaller levels of self-similar detail as it is probed in ever smaller scales (or equivalently, higher resolutions), represented here by A0. The scaling exponent between the measured intrinsic and extrinsic sizes is the so-called fractal dimension.

Revised version: This re-writing highlights a new perspective, or interpretation of the scaling law: it now suggests a relationship between intrinsic and extrinsic measures of cortical size2: At and Ae respectively, measured in units of A0. This is reminiscent of fractal scaling [9], where a complex shape reveals ever smaller levels of self-similar detail as it is probed in ever smaller scales (or equivalently, higher resolutions), represented here by A0. The scaling, or power exponent between the measured intrinsic and extrinsic sizes is the so-called fractal dimension.

Earlier version: Although actual fractals are mathematical abstractions, they can often be defined as the limit of iterative processes. Many structures in nature, and in particular biology [10, 11, 12, 13, 14, 15, 16], are good approximations of a fractal. Equation (2) is suggestive, but not proof, that cortices are among these forms, with a fractal dimension of 1.25 × 2 = 2.5 (the factor 2 is the topological dimension of areas). Indeed, fractal scaling for various aspects of cortical morphology has often been postulated [17, 18], with a number of recent papers making use of MRI data [19, 20, 21]. Most recently published estimates of fractal dimension for the whole cortex are indeed close to 2.5 [22, 23].

Revised version: Although actual fractals are mathematical abstractions, they can often be defined as the limit of iterative processes. Many structures in nature, and in particular biology [10, 11, 12, 13, 14, 15, 16], are good approximations of a fractal. Equation (2) is suggestive, but not proof, that cortices are among these forms, with a fractal dimension of 1.25 × 2 = 2.5 (the factor 2 being the topological dimension of areas). Indeed, fractal scaling for various aspects of cortical morphology has often been postulated [17, 18], with a number of recent papers making use of MRI data [19, 20, 21]. Most recently published estimates of fractal dimension for the whole cortex are indeed close to 2.5 [22, 23, 24].

Earlier version: Here, for the first time, we show how to directly construct morphologically plausible realisations of cortices at any specified spatial scale, or resolution. This is achieved through a coarse-graining procedure that removes morphological details smaller than a specified scale while preserving surface integrity. For example, at a set scale of 3 mm, sulcal walls that are less than 3 mm apart would be removed, and the neighbouring gyri could be fused. This method is a new systematic way of expressing cortical shape, in terms of a flow of morphometric measurements as spatial scale varies. By examining how areas scale across coarse-grained realisations of actual primate cortices, we will be able to directly verify cortical self-similarity.

Revised version: Here, for the first time, we propose to directly construct morphologically plausible realisations of cortices at any specified spatial scale, or resolution. This is achieved through a coarse-graining method that removes morphological details smaller than a specified scale while preserving surface integrity. For example, at a set scale of 3 mm, sulcal walls that are less than 3 mm apart would be removed, and the neighbouring gyri would be fused. This method is a new systematic way of expressing cortical shape, in terms of a flow of morphometric measurements as spatial scale varies. By examining how areas scale across coarse-grained realisations of actual primate cortices, we will be able to directly verify cortical self-similarity.

Added in the revised version: As a starting point for such a coarse-graining procedure, we suggest to turn to a well-established method that measures fractal dimension of objects: the so-called box-counting algorithm [25, 26]. Briefly, this algorithm fill the object of interest (say the cortex in our case) with boxes, or voxels of increasingly larger sizes and counts the number of boxes in the object as a function of box size. As the box size increases, the number of boxes decreases; and in a log-log plot, the slope of this relationship indicates the fractal dimension of the object. In our case, this method would not only provide us with the fractal dimension of the cortex, but, with increasing box size, the filled cortex would also contain less and less detail of the folded shape of the cortex. Intuitively, with increasing box size, the smaller details below the resolution of a single box would disappear first and increasingly larger details will follow – precisely what we require from a coarse-graining method. We therefore propose to expand the traditional box-counting method beyond its use to measure fractal dimension, but to also analyse the reconstructed cortices as different realisations of the original cortex at the specified spatial scale.

Earlier version: Our proposed method requires the bounding pial and white matter surfaces of the cortical ribbon as input. We obtained these surfaces based on reconstructions from magnetic resonance imaging data in 11 different primate species (see Suppl. S1). Algorithmically, we then segment the space between the original pial and white matter surfaces into a 3D grid of cubes of the desired scale λ, where each cube is of dimensions λ × λ × λ. We also term the 3D grid of cubes “voxelisation”, as it effectively captures the shape of the cerebral cortex as voxels in 3D space (Fig. 1B bottom row). At any given scale, or voxel size, this process effectively erases morphological features (folds) that are smaller than the cube size. Visually, increasing the voxel size appears as if the cortex is “melting” (Fig. 1, and videos: https://bit.ly/3CDoqZQ). A detailed description and discussion of the algorithm in provided in Suppl. S2.

Revised version: Concretely, our proposed method requires the bounding pial and white matter surfaces of the cortical ribbon as input. We obtained these surfaces based on reconstructions from magnetic resonance imaging data in 11 different primate species (see Suppl. S1). Algorithmically, we then segment the space between the original pial and white matter surfaces into a 3D grid of boxes of the desired scale λ, where each box is a cube of dimensions λ × λ × λ. We also term the 3D grid of cubes “voxelisation”, as it effectively captures the shape of the cerebral cortex as voxels in 3D space (Fig. 1B bottom row). At any given scale, or voxel size, this process effectively erases morphological features (folds) that are smaller than the cube size. Visually, increasing the voxel size appears as if the cortex is “melting” and “thickening” (Fig. 1, and videos: https://bit.ly/3CDoqZQ).

Added in the revised version: The observation that with increasing voxel sizes, the coarse-grained cortices tend to be smoother and thicker is particularly interesting: the scaling law in Eq. 3 can be understood as thicker cortices (T) form larger folds (or are smoother i.e. less surface area At) when brain size is kept constant (Ae). This way of understanding has also been vividly illustrated by using the analogy of forming paper balls with papers of varying thickness in Mota et al.[4]: to achieve the same size of a paper ball (Ae), the one that uses thicker paper (T) will show larger folds (or is smoother i.e. less surface area At) than the one using thinner paper. The scaling law can therefore be understood as a physically and biologically plausible statement, and here, we are encouraged that our algorithm yields results in line with the scaling law.

Added in the revised version: A more technical and detailed description and discussion of the algorithm is provided in Suppl. S2. Note this method has also no direct dependency on the original MR image resolution, as the inputs are smooth grey and white matter surface meshes reconstructed from the images using strong (bio-)physical assumptions and therefore containing more fine-grained spatial information than the raw images (also see Suppl. S3).

Earlier version: Finally, to analyse the scaling behaviour of the different coarse-grained realisations of the same brain, we apply an isometric rescaling process that leaves all dimensionless shape properties unaffected (more details in Suppl. S3.1). Conceptually, this process fixes the voxel size, and instead resizes the surfaces accordingly, which ensures that we can compare the coarse-grained realisations to the original cortices, and test if the former, like the latter, also scale as in Eqn. (1).

Revised version: Finally, to analyse the scaling behaviour of the different coarse-grained realisations of the same brain, we apply an isometric rescaling process that leaves all dimensionless shape properties unaffected (more details in Suppl. S4.1). Conceptually, this process fixes the voxel size, and instead resizes the surfaces relative to the voxel size, which ensures that we can compare the coarse-grained realisations to the original cortices, and test if the former, like the latter, also scale according to Eqn. (1). Resizing, or more precisely, shrinking the cortical surface is mathematically equivalent to increasing the box size in our coarse-graining method. Both achieved an erasure of folding details below a certain threshold.

Added in the revised version: Coarse-graining a cortex at different scales.

Earlier version: Coarse-graining a cortex at different scales. A: Example original pial surface from a healthy human viewed from varying angles. B: Same example cortical surface from A, coarse-grained to different spatial scales. Top row shows resulting pial surfaces (with a small amount of smoothing applied for visualisation purposes here). Bottom row shows the corresponding voxelisation at each scale through the slice indicated by blue plane in panel A.

Revised version: A: Example original pial surface from a healthy human viewed from varying angles. B: Same example cortical surface from A, coarse-grained to different spatial scales. Top row shows resulting pial surfaces (with a small amount of smoothing applied for visualisation purposes here). Bottom row shows the corresponding voxelisation at each scale through the slice indicated by blue plane in panel A.

3 Results

3.1 All primate brains follow the same scaling law across spatial scales

Earlier version: We have analysed cortices of 11 different primate species (see Suppl. S1): aotus, cebus, chimpanzee, colobus, galago, human, lagothrix, lophocebus, macaque, marmoset, and pithecia, and various cohorts of human subjects. We applied our coarse-graining procedure to their pial and white matter surfaces, and empirically determined (i) that all species followed a power law (linear regression R2 > 0.999 for all species); (ii) the slope of said power law is α = 1.255 on a group level (CI: [1.254 1.256]) using linear mixed effect modelling (see Fig. 2 for visualisation, and Suppl. S3 for a detailed breakdown by species); and importantly, (iii) all species also show a similar offset log k ≈ −0.65263, with a standard deviation of intercept across species estimated at 0.02 from linear mixed effect modelling.

Revised version: We have analysed cortices of 11 different primate species (see Suppl. S1): aotus, cebus, chimpanzee, colobus, galago, human, lagothrix, lophocebus, macaque, marmoset, and pithecia, and various cohorts of human subjects. We applied our coarse-graining procedure to their pial and white matter surfaces, and empirically determined (i) that all species followed a power law (linear regression R2 > 0.999 for all species); (ii) the slope of said power law is α = 1.255 on a group level (CI: [1.254 1.256]) using linear mixed effect modelling (see Fig. 2 for visualisation, and Suppl. S4 for a detailed breakdown by species); and importantly, (iii) all species also show a similar offset log k ≈ −0.65263, with a standard deviation of intercept across species estimated at 0.02 from linear mixed effect modelling in log k.

Earlier version: Taken separately, the scaling for each species is proof that their cortices are self-similar with the same scaling: they each approximate a fractal with fractal dimension df = 2.5. Considering all species together, different species also overlap substantially (similar offset), and only differ from each other in the range of scales over which the approximation is valid (see Suppl. S3.2 and S4). Thus, as Fig. 2 illustrates, the data supports a universal scaling law across primate species and spatial scales: At(λ)T(λ)12=kAe(λ)54 (3)

Revised version: Taken separately, the scaling for each species is proof that their cortices are self-similar with the same scaling: they each approximate a fractal with fractal dimension df = 2.5. Considering all species together, different species also overlap substantially (similar offset), and only differ from each other in the range of scales over which the approximation is valid (see Suppl. S4.2 and S5). Thus, as Fig. 2 illustrates, the data supports a universal scaling law across primate species and spatial scales: A t λT λ 1 2 =k A e λ 5 4 , (3) with k = 0.2277.

Removed in the revised version: with k = 0.2277.

Universal scaling law for 11 coarse-grained primate brains.

Earlier version: A: Coarse-grained primate brains are shown in terms of their relationship between log10(At×T) vs. log10(Ae). Each solid line indicates a cortical hemisphere from a primate species. Thin grey lines indicate a slope of 1 for reference. Filled circles mark the data points of the original cortical surfaces. Grey data points are plotted for reference and show the comparative neuroanatomy dataset across a range of mammalian brains [5]. B: Slopes (a) of the regression of data points in A for each species. C: Rescaled brain surfaces visualised for five example species at different levels of coarse-graining.

Revised version: A: Coarse-grained primate brains are shown in terms of their relationship between A t × T vs. log10(Ae). Each solid line indicates a cortical hemisphere from a primate species. Thin grey lines indicate a slope of 1 for reference. Filled circles mark the data points of the original cortical surfaces. Grey data points are plotted for reference and show the comparative neuroanatomy dataset across a range of mammalian brains [5]. B: Slopes (α) of the regression of data points in A for each species. For each species, two data points are shown, one per hemisphere. Colour-code for each species is maintained throughout the whole figure. C: Rescaled brain surfaces visualised for five example species at different levels of coarse-graining.

3.2 Primate brains at different spatial scales are morphometrically similar to each other and other mammalian species

Earlier version: To better characterise the coarse-grained cortices in terms of their similarity in offset, we use a previously introduced [24] set of independent morphometric measures, K, I and S, that summarize the shape of the cortex in a natural and statistically robust way. In this framework, isometrically scaled copies of the same shape all map onto a line along the I = log At + log Ae + log T2 direction, which is perpendicular to a K × S plane that fully summarizes their shape. K=logAt−54logAe+14logT2 is the direction defined by the offset k of the scaling law Eqn. (1), while direction S=32logAt+34logAe−94logT2 captures the remaining information about shape, and can be regarded as a simple measure of morphological complexity. Ref [24] provides a detailed derivation and demonstration of the superior sensitivity and specificity of these new morphometric measures. The advantage of using this framework here is that we can assess the offset K (and shape term S) without interference by isometric size effects.

Revised version: To better characterise the coarse-grained cortices in terms of their similarity in offset, we use a previously introduced [27] set of independent morphometric measures, K, I and S, that summarise the shape of the cortex in a natural and statistically robust way. In this framework, isometrically scaled copies of the same shape all map onto a line along the I = log At + log Ae + log T2 direction, which is perpendicular to a K × S plane that fully summarises their shape. K=log A t − 5 4 log A e + 1 4 log T 2 is the direction defined by the offset k of the scaling law Eqn. (1), while direction S= 3 2 log A t + 3 4 log A e − 9 4 log T 2 captures the remaining information about shape, and can be regarded as a simple measure of morphological complexity. Ref [27] provides a detailed derivation and demonstration of the superior sensitivity and specificity of these new morphometric measures. The advantage of using this framework here is that we can assess the offset K (and shape term S) without interference by isometric size effects, including any re-scaling procedures. In other words, the following results are independent of the isometric size of the brain surfaces analysed.

Earlier version: One can directly compute the values of K and S for any cortex-like object, thus ascribing to it a point in the K × S plane. But a fuller expression of a particular shape is captured by the trajectory of said object as a function of coarse-graining. This is a very convenient and informative way of summarising shape: Self-similar objects correspond to straight trajectories. In particular, objects without any folds or protrusion (i.e., convex, such as the box with finite thickness in Fig. 3 A) correspond to the line K=−19S , as Ae = At for all levels of coarse-graining; and to horizontal trajectories (K = constant) if df = 2.5 (Fig. 3 B). A shape is said to be universal for a set of objects when their trajectories overlap, so that they can all be regarded as coarse-grained versions of one another (Fig. 3 C).

Revised version: One can directly compute the values of K and S for any cortex-like object, thus ascribing to it a point in the K × S plane. But a fuller expression of a particular shape is captured by the trajectory of said object as a function of coarse-graining in the K × S plane. This is a very convenient and informative way of summarising shape: self-similar objects correspond to straight trajectories as the K × S plane is in log-log space. In particular, objects without any folds or protrusion (i.e., convex, such as the box with finite thickness in Fig. 3 A) correspond to the line K=− 1 9 S , as Ae = At for all levels of coarse-graining. Horizontal trajectories (constant K) represent fractal objects with fractal dimension df = 2.5 (Fig. 3 B). And finally, in the K × S plane, a shape is said to be “universal” for a set of objects when their trajectories overlap, so that they can all be regarded as coarse-grained versions of one another (Fig. 3 C).

From this perspective, the regularity observed in primate cortices is remarkable (Fig. 3 D). The value of K is nearly invariant in all cases, in parallel to what was found in [4]. But, over all levels of coarse-graining, K also remains near-invariant in all trajectories as S decreases, resulting in a set of horizontal lines that largely overlap with each other and other mammalian species. The variance in K across scales and all 11 species is < 0.01, which is at least an order of magnitude lower than the variance in S. Primate brains therefore have all three characteristics of self-similarity, fractality (with df = 2.5), and universality (invariant K for all scales and species) at the same time.

Earlier version: Thus, coarse-grained primate cortical shapes are morphometrically similar to and in terms of the universal law ‘as valid as’ actual existing mammalian cortices. In contrast, we tested various non-brain shapes, and while e.g. the walnut, and bell pepper shapes form (partially) straight lines, they vary in both K and S (see Suppl. S5). These shapes may have a fractal regime, but their fractal dimension is not 2.5, nor is their shape similar to primate or mammalian brains in terms of K or S. Furthermore, supplementary S6 underscores the algorithmic and statistical robustness of these results using multiple realisations of the coarse-graining procedure on the same shape. In this framework, our main result can thus be expressed simply: For all cortices we analyzed, and for none of the non-cortices, coarse-graining will leave K largely unaffected, while morphological complexity S will decrease.

Revised version: Thus, coarse-grained primate cortical shapes are morphometrically similar to, and in terms of the universal law, ‘as valid as’ actual existing mammalian cortices3. In contrast, we tested various non-brain shapes, and while e.g. the walnut, and bell pepper shapes form (partially) straight lines, they vary in both K and S (see Suppl. S6). These shapes may have a fractal regime, but their fractal dimension is not 2.5, nor is their shape similar to primate or mammalian brains in terms of K or S. Furthermore, supplementary S7 underscores the algorithmic and statistical robustness of these results using multiple realisations of the coarse-graining procedure on the same shape. In this framework, our main result can thus be expressed simply: for all the cortices we analysed, and for none of the non-cortices, coarse-graining will leave K largely unaffected, while morphological complexity S will decrease.

Added in the revised version: 3.3 Morphometric measures as functions of scale reveal scale-specific effects of ageing

Added in the revised version: In the final part of our work, we show how our algorithm and the associated new understanding of brain morphology may become useful in applications. As an example, we will focus on how the ageing process affects human cortical morphology across scales. In Fig. 4 A, we compare the total surface area At(λ) as a function of scale λ for a young (20 year old) vs. an old (80 year old) group of human brains. Strikingly, the difference in At between the groups takes a U shape, where the strongest effect is seen at approx 2 millimeter (older subjects have higher At). For scales over approx 5 mm, and under 0.5 mm the difference become relatively small, suggesting the ageing process has less effect on the largest and smallest cortical morphological features. Finally, we reproduced these results in an independent dataset in Suppl. S8.2.

Added in the revised version: In this particular example, the dramatic scale-dependency of morphological measures can be visually and intuitively understood by looking at the reconstructed surfaces at each scale: Fig. 4 B shows some coronal slices of the cortical surface. At scale 0.27 mm, the gyri in the younger subjects are densely packed, but the older subjects show the expected widening between gyral walls and decrease in gyral surface area at the crown (see e.g. [28, 26] for recent investigations and references therein). At 1.86 mm, the younger cortices have already partially “melted”, erasing most small sulci between the densely packed gyri. In the older humans, however, the gyri are less dense, the sulci more open, and thus, at 1.86 mm, most gyri and sulci have not been erased yet. At scale 7.94 mm, both young and old brains have “melted” down to similarly near-lissencephalic shapes. For a more detailed multiscale investigation over the entire human lifespan, please refer to our new preprint [?]

Trajectories of coarse-grained primate cortices and other mammalian and human brains in K × S plane.

Earlier version: A: Straight trajectories indicate self-similarity (described by a scaling law). In particular, the black line here indicates objects with Ae = At for all scales, such as the box of finite thickness with a fractal dimension df = 1 (grey data points). This line is reproduced in all subpanels for reference. B: Multiple hypothetical flat trajectories are shown which correspond to df = 2.5 in this space. C: Hypothetical overlapping straight trajectories indicate multiple realisations of the same fractal object. D: Projecting our actual data into the normalised K × S plane showing the coarse-grained primate brains (same as in Fig. 2) as data point connected with solid lines. Different mammalian brains are shown as grey scatter points, and adult human data points are blue.

Revised version: A: Straight trajectories indicate self-similarity (described by a scaling law). In particular, the black line here indicates objects with Ae = At for all scales, such as the box of finite thickness with a fractal dimension df = 1 (grey data points). This line is reproduced in all subpanels for reference. B: Morphological trajectories of multiple hypothetical fractal objects are shown which. A flat trajectory (constant K) correspond to df = 2.5 in this space. However, these objects are clearly different fractals with different values of K. C: Hypothetical objects with overlapping straight morphological trajectories indicate multiple realisations of the same fractal object. Flat trajectories (constant K) correspond to df = 2.5. The two hypothetical objects with a decreasing K(S) correspond to 2 < df < 2.5 D: Projecting our actual data into the normalised K × S plane showing the coarse-grained primate brains (same as in Fig. 2) as data point connected with solid lines (colour-code same as Fig. 2). Different mammalian brains are shown as grey scatter points, and adult human data points are blue.

Removed in the revised version: 3.3 Morphometric measures as functions of scale reveal scale-specific effects of ageing

Removed in the revised version: In the final part of our work, we show how our algorithm and the associated new understanding of brain morphology may become useful in applications. As an example, we will focus on how the ageing process affects human cortical morphology across scales. In Fig. 4 A, we compare K(λ) as a function of scale λ for a young (20 year old) vs. an old (80 year old) group of human brains. Strikingly, the difference in K between the groups takes a U shape, where the strongest positive effect is seen at sub-millimeter scales, but the effect reverses and turns negative at millimeter scales (older subjects have higher K), peaking at about 2 mm. For scales over 5 mm the differences become negligible, suggesting the aging process has little effect on the larger cortical morphological features. Finally, we reproduced these results in an independent dataset in Suppl. S7.2.

Removed in the revised version: In this particular example, the dramatic scale-dependency of the effect of K can be visually and intuitively understood by looking at changes in total surface area (Fig. 4 B, also see Suppl. S7.1): Fig. 4 C shows some coronal slices of the cortical surface. At scale 0.27 mm, the gyri in the younger subjects are densely packed, but the older subjects show the expected (see e.g. [25, 26] for recent investigations and references therein) widening between gyral walls and decrease in gyral surface area at the crown. At 1.86 mm, the younger cortices have already partially “melted”, erasing most small sulci between the densely packed gyri. In the older humans, however, the gyri are less dense, the sulci more open, and thus, at 1.86 mm, most gyri and sulci have not been erased yet. At scale 7.94 mm, both young and old brains have “melted” down to very similar near-lissencephalic shapes.

More broadly, one can regard the melting process as a way of determining how cortical area is allocated across different scales. As the cortex ‘melts’, the contributions to the total area from features smaller than the cut-off scale are eliminated. In the example in Fig. 4, for instance, we can say about half (105−4.7 = 100.3 ≈ 2) the total area in the 80 y.o. cortices is present in features smaller than 4 mm.

Added in the revised version: From an application perspective, this example here illustrates clearly that the surface area difference between older and younger subjects at “native” scale (i.e. original freesurfer surfaces) is negligible (effect size smaller than 2 standard deviations). However, in our analysis across scales, there is a clear optimal scale at approx. 2 mm where the effect size is maximised between older and younger subjects (effect size is approx-8 standard deviations). For most classification applications, the scale-dependent morphometrics therefore represent a huge advance over the native scale.

Human ageing shows differential effects depending on spatial scale.

Earlier version: A: Top: K(λ) is shown for a group of 20-year-olds (red) and a group of 80-year-olds (blue). Individual data points show individual subjects. Mean and standard deviation are shown as the solid line and the shaded area respectively. Bottom: Effect size (measured as ranksum z-values) between the older and younger groups at each scale. Positive effect indicates a larger value for the younger group. Blue arrows indicate the effect size at “native scale” i.e. using the original grey and white matter meshes. B: Same as A, but for log10 At (without rescaling). C: Coronal slices of the pial surface of a 20-year-old and an 80-year-old at different scales (columns).

Revised version: A: Top: At (λ) is shown for a group of 20-year-olds (red, n=27) and a group of 80-year-olds (blue, n=86). Mean and standard deviation are shown as the solid line and the shaded area respectively. Bottomml: Effect size (measured as ranksum z-values) between the older and younger groups at each scale. Positive effect indicates a larger value for the younger group. Blue arrows indicate the effect size at “native scale” i.e. using the original freesurfer grey and white matter meshes. B: Coronal slices of the pial surface of an example 20-year-old subject and an 80-year-old subject at different scales (columns).

4 Discussion

Earlier version: We have devised a new way of expressing the shape of the mammalian cerebral cortex, as the flow in the values of morphometric measures over a range of spatial scales. This was achieved by coarse-graining cortical surfaces, erasing morphological features smaller than the specified scale, while preserving surface integrity. After applying this method to the cortices of 11 primate species, we have shown that all these diverse shapes are all approximations of same archetypal self-similar (fractal) shape. Most of their morphological diversity can be ascribed to the species-specific ranges of spatial scales over which each approximation is valid, with the smallest scale being an invariant multiple of cortical thickness. This was a proof-by-construction of fractality, a step beyond the usual box-counting approach, which also yields scale-dependent morphometrics. As a proof-of-principle we showed that healthy human ageing has highly scale-dependent effects in a range of morphometrics.

Revised version: We have devised a new way of expressing the shape of the mammalian cerebral cortex, as the flow in the values of morphometric measures over a range of spatial scales. This was achieved by coarse-graining cortical surfaces, erasing morphological features smaller than the specified scale, while preserving surface integrity. After applying this method to the cortices of 11 primate species, we have shown that all these diverse shapes are approximations of the same archetypal self-similar (fractal) shape. Most of their morphological diversity can be ascribed to the species-specific ranges of spatial scales over which each approximation is valid, with the smallest scale being an invariant multiple of cortical thickness. This was a proof-by-construction of fractality, a step beyond the usual box-counting approach, which also yields scale-dependent morphometrics. As a proof-of- principle we showed that healthy human ageing has highly scale-dependent effects in a range of morphometrics.

Added in the revised version: Compared to previous literature, we can summaries our main contribution and advance as follows: (i) We are showing for the first time that representative primate species follow the exact same fractal scaling – as opposed to previous work showing that they have a similar fractal dimension [29, 30], i.e. slope, but not necessarily the same offset, as previous methods had no consistent way of comparing offsets. (ii) Previous work could also not show direct agreement in morphometrics between the coarse-grained brains of primate species and other non-primate mammalian species. (iii) Demonstrating in proof-of-principle that multiscale morphometrics, in practice, can have much larger effect sizes for classification applications. This moves beyond our previous work where we only showed the scaling law across [4] and within species [6], but all on one (native) scale with comparable effect sizes for classification applications [6].

Added in the revised version: In simple terms: we know that objects can have the same fractal dimension, but differ greatly in a range of other shape properties. However, we demonstrate here, that representative primate brains and mammalian brain indeed share a range of other key shape properties, on top of agreeing in fractal dimension. This suggests a universal blueprint for mammalian brain shape and a common set of mechanisms governing cortical folding. As a practical additional outcome of our study, we could show that our novel method of deriving multiscale metrics of brain shape can differentiate subtle shape changes much better than the metrics we have been using so far at a single native scale.

Earlier version: Expressing cortical morphology as a function of scale is more detailed than a list of summary morphometric measures, and more informative than the mere listing of every sulci and gyri. We propose this new syntax as the basis for a more rigorous characterization of brain morphology and morphological changes. A clear advantage is that some biological processes may only act on a specific spatial scale, leaving other scales untouched (ageing in our example). By disambiguation of the spatial scale, it allows for an extra dimension of understanding, and tracking of biological processes. In the future, this approach can also be extended to cortical development and to various degenerative [6] and congenital [24] neuropathic conditions, especially if combined with a regionalized version of this method applied to specific cortical regions [7] or local patches [8].

Revised version: Expressing cortical morphology as a function of scale is more detailed than a list of summary morphometric measures, and more informative than the mere listing of every sulcus and gyrus. We propose this new syntax as the basis for a more rigorous characterisation of brain morphology and morphological changes. A clear advantage is that some biological processes may only act on a specific spatial scale, leaving other scales untouched (ageing in our example). By disambiguation of the spatial scale, it allows for an extra dimension of understanding, and tracking of biological processes. In the future, this approach can also be extended to cortical development and to various degenerative [6] and congenital [27] neuropathic conditions, especially if combined with a regionalised version of this method applied to specific cortical regions [7] or local patches [8].

Earlier version: Empirically, the main result of this paper is the demonstration of a universal self-similar scaling (Eqn. 3) for primate, and presumably mammalian cortices. There are two aspects of this universality: First that for each and every cortex the value of K remains the same for all scales as one removes substructures smaller than a varying length scale (or equivalently, that the fractal dimension is almost exactly 2.5 in all cases). Second, that the value of K for the cortices of different species is approximately the same, as previously observed (Eqn. 1) across species [5] and individuals [6]. One could imagine a set of shapes for which one but not the other aspect of the universality in K holds true. However, the fact that both universalities hold true is significant. It suggests the existence of a single highly conserved mechanism for cortical folding, operating on all length scales self-similarly with only a few morphological degrees of freedom. It also hints at the possibility of deriving cortical scaling from some variational principle. Finally, this dual universality is also a more stringent test for existing and future models of cortical gyrification mechanisms, and one that moreover is applicable to individual cortices.

Revised version: Empirically, the main result of this paper is the demonstration of a universal self-similar scaling (Eqn. 3) for primate, and presumably mammalian cortices. There are two aspects of this universality: first that for each and every cortex the value of K remains the same for all scales as one removes substructures smaller than a varying length scale (or equivalently, that the fractal dimension is almost exactly 2.5 in all cases). Second, that the value of K for the cortices of different species is approximately the same, as previously observed (Eqn. 1) across species [5] and individuals [6]. One could imagine a set of shapes for which one but not the other aspect of the universality in K holds true. However, the fact that both universalities hold true is significant. It suggests the existence of a single highly conserved mechanism for cortical folding, operating on all length scales self-similarly with only a few morphological degrees of freedom. It also hints at the possibility of deriving cortical scaling from some variational principle. Finally, this dual universality is also a more stringent test for existing and future models of cortical gyrification mechanisms at relevant scales, and one that moreover is applicable to individual cortices. For example, any models that explicitly simulate a cortical surface could be directly coarse-grained with our method and compared to actual human and primate data provided here.

Earlier version: The scaling itself does not imply or favour any particular proposed gyrification model (ours [5] included), and all results in this paper are agnostic about this choice. Indeed, our previously proposed model [5] for cortical gyrification is very simple, which is both a strength and a weakness. There are many complex phenomena in nature that display simple and universal scaling that can be derived from first principles [27, 28, 29]. In addition, recent results suggest simplicity and symmetry are generically favoured on statistical-ensemble grounds by evolution [30]. Our model correctly predicts the scaling law (Eqn. 1), but a more complete explanation for cortical gyrification is probably far more complex [31] than can be accounted by such a simple model.

Revised version: The scaling itself does not imply or favour any particular proposed gyrification model (ours [5] included), and all results in this paper are agnostic about this choice. Indeed, our previously proposed model [5] for cortical gyrification is very simple, assuming only a self-avoiding cortex of finite thickness experiencing pressures (e.g. exerted by white matter pulling, or by CSF pressure). The offset K, or ‘tension term’, precisely relates to these pressures, leading us to speculate that subtle changes in K correlate with changes in white matter property [6, 27]. In the same vein of speculation, the scale-dependence of K shown in this work might therefore be related to different types of white matter that span different length scales, such as superficial vs. deep white matter, or U-fibres vs. major tracts. However, there are also challenges to the axonal tension hypothesis [31]. Indeed, white matter tension differentials in the developed brain may not explain location of folds, but instead white matter tension may contribute to a whole-brain scale ‘pressure’ during development that drives the folding process overall. Aside from speculations about the biological interpretation, the simplicity of the highlighted scaling law parallels many complex phenomena in nature that display simple and universal scaling that can be derived from first principles [32, 33, 34]. In addition, recent results suggest simplicity and symmetry are generically favoured on statistical- ensemble grounds by evolution [35]. Our model correctly predicts the scaling law (Eqn. 1), but a more complete explanation for cortical gyrification is probably far more complex [36] than can be accounted by such a simple model.

Earlier version: One specific example of said complexity are the exact patterns, locations, depth, and features of gyri and sulci. We know such patterns to be, for example, somewhat heritable in humans, whilst in macaques such patterns are relatively preserved across the species. Our work does not explain any of these observations, nor are the coarse-grained versions of human brains supposed to exactly resemble the location/pattern/features of gyri and sulci of other primates. The similarity we highlighted here are on the level of summary metrics, and our goal was to highlight the universality in such metrics to point towards highly conserved quantities and mechanisms.

Revised version: One specific example of said complexity are the exact patterns, locations, depth, and features of gyri and sulci. We know such patterns to be, for example, variable but also somewhat heritable in humans, whilst in macaques such patterns are relatively preserved across the species. Our work does not explain any of these observations, nor are the coarse-grained versions of human brains supposed to exactly resemble the location/pattern/features of gyri and sulci of other primates. The similarity we highlighted here are on the level of summary metrics, and our goal was to highlight the universality in such metrics to point towards highly conserved quantities and mechanisms.

Earlier version: More broadly, the interaction between brain development and evolution may also benefit from a scale-specific understanding. This may be important in elucidating what in cortical morphology is selected for by evolution, what is determined by physics; what is specified by genes, and what is emergent. For example, one can estimate the number of structural features at each scale (At(λ) as multiples of A0(λ), and it will be interesting to correlate this number to other quantifiers of cortical structure, such as neuron numbers [4], number of functional areas [32] or the number of cortical columns [33], possibly over different stages of development. Generally, the larger the cortical feature (i.e. from gyri to functional areas to lobes to hemispheres), the earlier during development it appears [34, 35], and more broadly it is conserved over kinship [36] and phylogeny [37, 38]. It thus seems likely that comparative neuroanatomical methods [39, 40] may be directly used to identify and contrast structures in coarse-grained cortices of more highly gyrified species with their analogues in less-gyrified species. One could then perhaps specify when evolution conserves and when it invents old and new cortical structures.

Revised version: More broadly, the interaction between brain development and evolution may also benefit from a scale-specific understanding. This may be important in elucidating what in cortical morphology is selected for by evolution, what is determined by physics; what is specified by genes, and what is emergent. For example, one can estimate the number of structural features at each scale (At(λ) as multiples of A0(λ)), and it will be interesting to correlate this number to other quantifiers of cortical structure, such as number of neurons [4], number of functional areas [37] or the number of cortical columns [38], possibly over different stages of development. Generally, the larger the cortical feature (i.e. from gyri to functional areas to lobes to hemispheres), the earlier during development it appears [39, 40], and more broadly it is conserved over kinship [41] and phylogeny [42, 43]. It thus seems likely that comparative neuroanatomical methods [44, 45] may be directly used to identify and contrast structures in coarse-grained cortices of more highly gyrified species with their analogues in less-gyrified species. One could then perhaps specify when evolution conserves and when it invents old and new cortical features.

Earlier version: Our work here was limited to summary descriptors of entire cortical hemispheres, but future work will explore extensions of these methods to lobes and cortical areas, similarly to [7, 8]. This will generate precise characterizations of the morphological differences between phylae and across developmental stages, and perhaps pinpoint the time and location of morphological changes leading to congenital and neurodegenerative conditions. Ultimately, we hope this new framework for expressing and analyzing cortical morphology, besides revealing a hitherto hidden regularity of nature, can become a powerful tool to characterize and compare cortices of different species and individuals, across development and aging, and across health and disease.

Revised version: Our work here was limited to summary descriptors of entire cortical hemispheres, but future work will explore extensions of these methods to lobes and cortical areas, similarly to [7, 8]. This will generate precise characterisations of the morphological differences between phylae and across developmental stages, and perhaps pinpoint the time and location of morphological changes leading to congenital and neurodegenerative conditions. Ultimately, we hope this new framework for expressing and analysing cortical morphology, besides revealing a hitherto hidden regularity of nature, can become a powerful tool to characterise and compare cortices of different species and individuals, across development and ageing, and across health and disease.